# Module 8 Assignment

A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
% matplotlib inline

import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt

from datetime import datetime, date, timedelta

from nose.tools import assert_equal, assert_true
```

-----

## Problem 1: Get Future Date

In the code cell below, we declare a function named `get_future_date` that takes two argument: `current_date`, which is a datetime object, and `days` which is an integer. The function will return the datetime object `days` after `current_date`.

To complete this problem, finish writing the function `get_future_date`:
- Add `days` into `current_date`(Hint: use `timedelta`)
- return the new datetime object.

-----

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
def get_future_date(current_date, days):
    # YOUR CODE HERE
    
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_equal(get_future_date(datetime(2019, 9,9), 1), datetime(2019, 9, 10), msg="Function is not correct." )
assert_equal(get_future_date(datetime(2019, 9,9), 366), datetime(2020, 9, 9), msg="Function is not correct." )
```

-----

## Problem 2: Reading in Data

In the code cell below, we declare a function named `read_data` that takes one function parameter: `file_path`, which is a string.

To complete this problem, finish writing the function `read_data`:
- read data from the file specified in the string `file_path`, by using the Pandas `read_csv` function, and
- return the resulting DataFrame.

-----

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
def read_data(file_path):
    '''
    Parameters
    ----------
    file_path: string containing path to the dataset
    
    Returns
    -------
    Pandas DataFrame
    '''
    # YOUR CODE HERE
    
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
path = 'data/dow_jones_index.data'
dow_df = read_data(path)
assert_equal(dow_df.shape[1], 16, msg="The number of columns your dataset has, does not match the solutions")
assert_equal(dow_df.shape[0], 750, msg="The number of rows your dataset has, does not match the solutions")
dow_df.head(2)
```

-----

## Problem 3: Set Datetime Index


For this problem you will use `dow_df` created in problem 1 autograder cell.

To solve this problem do the following:
- convert the date column in `dow_df` from `object` type to `datetime` type by using the Pandas function `to_datetime`,
- create datetime index from date column using pandas `DatetimeIndex` function and set it as the index of `dow_df`.

After this problem, dow_df has a datetime index.


------

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(dow_df.index), pd.core.indexes.datetimes.DatetimeIndex, msg="The index is not datetime object.")
dow_df.head(2)
```

---

# Problem 4: Slice DataFrame

Slice June 2011 stock information for 'AA'.

For this problem you will use **dow_df** updated in problem 2.

To solve this problem do the following:
- Use boolean mask to get rows of stock 'AA' from dow_df, assign the resulting DataFrame to aa_df
- Get all June 2011 data from aa_df and assign the resulting DataFrame to **aa_201106_df**

After this problem, there's a new variable **aa_201106_df** defined.

-----

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_equal(aa_201106_df.shape[0], 4, msg="aa_201106_df is not correct")
aa_201106_df
```

---

# Problem 5: Resample DataFrame

Resample stock 'AA'.

For this problem you will use **dow_df** updated in problem 2.

To solve this problem do the following:
- Use boolean mask to get rows of stock `AA` from dow_df, assign the resulting DataFrame to aa_df
- Resample aa_df with month end frequcy(Hint: with code 'M')
- Apply mean() function on the Resampler object and assign the resulting DataFrame to **aa_resample**

After this problem, there's a new variable **aa_resample** defined.


-----

<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(aa_resample), pd.DataFrame, msg="aa_resample should be a DataFrame")
assert_equal(aa_resample.shape[0], 6, msg="aa_resample is not correct")
assert_equal(aa_resample.index[0].date(), date(2011, 1,31),
             msg="aa_resample is not resampled at month end frequency")
aa_resample
```