# Module 8 Assignment

A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
%matplotlib inline

import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt

from datetime import datetime

from nose.tools import assert_equal, assert_true
```

-----

## Problem 1: Converting Datetme String to Datetime Object

In the Code cell below, we declare a function named `convert_datetime` that takes one argument: `datetime_str`, which is a string representation of datetime. The `datetime_str` is in "yyyy-mm-dd HH:MM:SS" format(ie. "2019-09-30 12:30:59")

To complete this problem, finish writing the function `convert_datetime`:
- Convert a datetime string in "yyyy-mm-dd HH:MM:SS" format to a datetime object.
- Return the datetime object.

-----

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
def convert_datetime(datetime_str):
    # YOUR CODE HERE
    
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_equal(convert_datetime('2019-09-09 00:00:00'), datetime(2019, 9, 9),
             msg='convert_datetime is not defined correctly')
assert_equal(convert_datetime('2019-09-09 12:59:59'), datetime(2019, 9, 9,12,59,59),
             msg='convert_datetime is not defined correctly')

```

-----

## Problem 2: Reading in Data

In the Code cell below, we declare a function named `read_data` that takes one function parameter: `file_path`, which is a string.

To complete this problem, finish writing the function `read_data`:
- Read data from the file specified in the string `file_path`, by using the Pandas `read_csv` function.
- Return the resulting DataFrame.

-----

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
def read_data(file_path):
    '''
    Parameters
    ----------
    file_path: string containing path to the dataset
    
    Returns
    -------
    Pandas DataFrame
    '''
    # YOUR CODE HERE
    
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
path = 'data/dow_jones_index.data'
dow_df = read_data(path)
assert_equal(dow_df.shape[1], 16, msg="The number of columns your dataset has, does not match the solutions")
assert_equal(dow_df.shape[0], 750, msg="The number of rows your dataset has, does not match the solutions")
dow_df.head(2)
```

-----

## Problem 3: Set Datetime Index


For this problem you will use `dow_df` created in problem 2 autograder cell.

To solve this problem do the following:
- Convert the date column in `dow_df` from `object` type to `datetime` type by using the Pandas function `to_datetime`.
- Create datetime index from date column using pandas `DatetimeIndex` function and set it as the index of `dow_df`.

After this problem, `dow_df` has a datetime index.

------

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(dow_df.index), pd.core.indexes.datetimes.DatetimeIndex, msg="The index is not datetime object.")
dow_df.head(2)
```

-----

## Problem 4: Plot AA Close Price

Plot close price of stock 'AA'.

For this problem you will use `dow_df` updated in problem 3.

To solve this problem do the following:
- Use boolean mask to get rows of stock `AA` from dow_df, assign the resulting DataFrame to **aa_df**
- Plot close price of aa_df by using the Pandas `plot` function.
- Set plot title to "AA close", xlabel to "date", ylabel to "close"

After this problem, a new variable **aa_df** is defined, a plot should show up with proper configuration.

-----

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
fig, ax = plt.subplots(figsize=(8, 4))

# YOUR CODE HERE

```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_equal(aa_df.shape[0], 25, msg='aa_df is not correct.')
assert_equal(ax.get_title(), 'AA close', msg='Plot title is not AA close')
assert_equal(ax.get_xlabel(), 'date', msg='Plot xlabel is not date')
assert_equal(ax.get_ylabel(), 'close', msg='Plot ylabel is not close')
```