# Module 2 Assignment

A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
import pandas as pd
import numpy as np
import seaborn as sns

from nose.tools import assert_equal, assert_almost_equal, assert_true, assert_is_instance
```

-----

# Predicting MPG

In this assignment, we will use mpg dataset to make a regression model. Before we attempt to build a mode, we first must load the data from the seaborn module.

Please run the next code cell before proceeding to Problem 1.

-----

```python
#load MPG dataset
mpg = sns.load_dataset('mpg')
mpg.head()
```

---

# Problem 1: Create Dependent and Independent Variables

For this problem you will use the DataFrame **mpg** defined above.

To complete the task, do the following:
1. Use `dmatrices` in `patsy` module to create independent variable **x** and dependent variable __y__ from DataFrame mpg. 
2. In patsy formula, use 'mpg' as dependent variable.
3. Use displacement, acceleration and model_year as independent varialbe.
4. Treat model_year as categorical feature(enclose in C()). 
5. Set `dmatrices` argument `return_type` to 'dataframe'.

After this problem, there are two new variables defined, **x** and __y__.

-----

<!-- grade_id: p1-ans, points: 0, grade: false, solution: true -->
```python
import patsy as pts

# YOUR CODE HERE

```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(x), pd.core.frame.DataFrame, msg="x is not a DataFrame")
assert_equal(x.shape, (398, 15), 
             msg="The independent features are not correct. Make sure to treat model_year as categorical feature")
# 2 random rows of the independent variables
x.sample(2)
```

-----

# Problem 2: Create the Training and Testing Datasets

This problem works on the variables **x** and __y__ created in problem 1.

Splite the independent and dependent variables to training and testing set.

To complete this process, do the following:

- Name the training and testing independent variable as x_train and x_test.
- Name the training and testing dependent variable as y_train and y_test.
- The `test_size` argument in `train_test_split` should be set to 0.3.
- The `random_state` argument in `train_test_split` should be set to 23.

After this problem, there are 4 new variables defined, **x_train, x_test, y_train** and __y_test__.

-----

<!-- grade_id: p2-ans, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# YOUR CODE HERE

```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
assert_equal(x_test.shape, (120, 15), 
             msg="The testing set size is not correct")
assert_equal(y_test.values[0][0], 16.0,
             msg="The testing dependent variable is not correct. Make sure you set random_state to 23")
#2 random rows of the training independent variables
x_train.sample(2)
```

---

# Problem 3: Performing Linear Regression

Your task for this problem is to build and use the scikit learn library's `LinearRegression` estimator to  make predictions on the mpg dataset. To complete this function, you must explicitly:
- Create a `LinearRegression` estimator **lin_model** by using scikit learn. Accept default values for all arguments.
- Fit the `LinearRegression` estimator using x_train and y_train created in problem 2.

After this problem, there will be a trained linear regression model **lin_model**.

-----

<!-- grade_id: p3-ans, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

# YOUR CODE HERE

```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(lin_model), type(LinearRegression()), msg="lin_model is not defined as a LinearRegression model")
assert_equal(lin_model.get_params()['fit_intercept'], True,
            msg="lin_model is not created with all default argument values")
```

---

# Problem 4: Checking R2 Score on Testing Dataset

For this problem, you will compute the R2 score of the lin_model.
To complete this function, you must explicitly:
- Compute the R2 score using `score` function of lin_model.
- Use x_test and y_test to calculate the score.
- Assign score to variable **r2_score**.

After this problem, there will be a new variable r2_score defined.

-----

<!-- grade_id: p4-ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(r2_score, 0.7325919749314689, msg="R2 score is not correct")

```