# Module 5 Assignment

A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
import pandas as pd
import numpy as np
import seaborn as sns

from nose.tools import assert_equal, assert_almost_equal, assert_true, assert_is_instance

# We do this to ignore warnings
import warnings
warnings.filterwarnings("ignore")
```

---
# Prepare Breast Cancer Data

This assignment will use the breast cancer dataset. Before we attempt to build a model, we first prepare the data.

Please run the next code cell before proceeding to Problem 1.


```python
#Load breast cancer dataset
df = pd.read_csv('data/breast-cancer-wisconsin.csv')
label = df['class']
data = df[['clump thickness', 'uniformity cell size', 'uniformity cell shape', 'marginal adhesion', 'epithelial cell size', 'bare nuclei', 'bland chromatin', 'normal nucleoli', 'mitoses']]
data.sample(2)
```

---
# Problem 1: Get Feature Variances by Variance Thresholding

Normalize features and get feature variances.

For this problem you will use the DataFrame **data** defined above.

To solve this problem do the following:
- Use `MinMaxScaler` to normalize DataFrame **data** created above, assign normalized data to variable __data_ss__.
- Create a `VarianceThreshold` feature selector with default `threshold`(which is 0) using scikit learns library.
- Fit & Transform the selector on the DataFrame **data_ss**.
- Retrieve feature variances from the selector's `variances_` attribute and assign it to variable **feature_variance**

After this problem, there's a new variable **feature_variance** defined.

---

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.preprocessing import MinMaxScaler
from sklearn.feature_selection import VarianceThreshold

# YOUR CODE HERE

```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(feature_variance[0], 0.09808697, msg='Feature variances are not correct')
assert_almost_equal(feature_variance[2], 0.11010541, msg='Feature variances are not correct')
print('Feature Variances:\n')
for feature, var in sorted(zip(data.columns, feature_variance), key=lambda x: x[1], reverse=True):
    print(f'{feature:>21} = {var:5.3f}')
```

---
# Problem 2: Get Feature Ranking by Recursive Feature Extraction

Perform RFE on a Random Forest Classifier and retrieve feature rankings.

For this problem you will use **data** and __label__ created above.

To solve this problem do the following:
- Create a `RandomForestClassifier` estimator. Set `n_estimators` to 100, `random_state` to 23 and accept default values for all other hyperparameters.
- Create a Recursive Feature Estimator `RFE` using the Random Forest Classifier created in step 1 as the `estimator`, set `n_features_to_select ` to 1. Accept default values for other arguments.
- Fit the RFE estimator using **data** and __lable__.
- Retrieve feature rankings from the `RFE` selector's `ranking_` attribute and assign it to variable **feature_ranking**.

After this problem, there's a new variable **feature_ranking** defined.

---

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier
from sklearn.feature_selection import RFE

# YOUR CODE HERE

```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
assert_equal(feature_ranking.tolist(), [6, 2, 1, 8, 5, 3, 4, 7, 9])
# Display feature ranking
print('Feature Ranking:')
for var, name in sorted(zip(feature_ranking, data.columns), key=lambda x: x[0]):
    print(f'{name:>21} = {var}')
```

---

# Problem 3: Get Feature Importance from Random Forest Classifier

Get feature importance from a trained Random Forest Classifier.

For this problem you will use **data** and __label__ created above.

To solve this problem do the following:
- Create a `RandomForestClassifier` estimator. Set `n_estimators` to 100, `random_state` to 23 and accept default values for all other hyperparameters.
- Fit the `RandomForestClassifier` estimator using **data** and __label__.
- Retrieve feature importances from the estimator's `feature_importances_` attribute and assign it to variable **feature_importance**.

After this problem, there will be a new variable **feature_importance** defined.

-----

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier

# YOUR CODE HERE

```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(feature_importance[0], 0.04282269, msg="Feature importances are not correct")
assert_almost_equal(feature_importance[-1], 0.00794928, msg="Feature importances are not correct")
print("Feature Importance:\n")
for val, name in sorted(zip(feature_importance, data.columns), key=lambda x: x[0], reverse=True):
    print(f'{name:>21}: {100.0*val:05.2f}%')
```

---

# Problem 4: Get the Cross Validation Scores

Get the cross-validation scores for a Random Forest Classifier.

For this problem you will use **data** and __label__ created above.

To solve this problem do the following:
- Create a `RandomForestClassifier` estimator. Set `n_estimators` to 100, `random_state` to 23 and accept default values for all other hyperparameters.
- Create `StratifiedKFold` iterator. Set `n_splits` to 5 and `random_state` to 23.
- Calculate cross validation scores using `cross_val_score` function with the random forest classifier, data, label and the `StratifiedKFold` iterator. Assign scores to variable **cv_scores**.

After this problem, there's a new variable **cv_scores** defined.

-----

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestClassifier
from sklearn.model_selection import StratifiedKFold
from sklearn.model_selection import cross_val_score

# YOUR CODE HERE

```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(cv_scores[0], 0.93430657, msg='Cross validation scores are not correct')
assert_almost_equal(cv_scores[2], 0.98540146, msg='Cross validation scores are not correct')
print(f"Average Cross Validation Score: {np.mean(cv_scores)*100:4.1f}%")
```